# Assignment Template

**Total Points: 10**

This is a template for creating Stata assignments with auto-grading.

```stata
* Setup - DO NOT MODIFY
clear all
sysuse auto, clear
```

Output:
```
(1978 automobile data)
```

## Question 1: Basic Calculation (2 points)

Calculate the mean of `price` and store it in a scalar called `mean_price`.

<!-- grade_id: q1_answer, points: 0, grade: false, solution: true -->
```stata
* YOUR CODE HERE
summarize price
scalar mean_price = r(mean)
```

Output:
```
Variable |        Obs        Mean    Std. dev.       Min        Max
-------------+---------------------------------------------------------
       price |         74    6165.257    2949.496       3291      15906
```

<!-- grade_id: q1_test, points: 2, grade: true, solution: false -->
```stata
* TEST CELL - DO NOT MODIFY
capture scalar list mean_price
assert _rc == 0

quietly summarize price
assert abs(mean_price - r(mean)) < 0.01

display "Test passed!"
```

Output:
```
Test passed!
```