# Simple Stata Quiz

**Total Points: 5**

This is a simpler example showing basic NBGrader patterns with Stata.

```stata
* Setup - DO NOT MODIFY
clear all
sysuse auto, clear
```

## Question 1: Calculate Mean (1 point)

Calculate the mean of `price` and store it in a scalar called `mean_price`.

<!-- grade_id: q1_answer, points: 0, grade: false, solution: true -->
```stata
* Calculate mean price
* YOUR CODE HERE
summarize price
scalar mean_price = r(mean)
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```stata
* Test
capture scalar list mean_price
assert _rc == 0
quietly summarize price
assert abs(mean_price - r(mean)) < 0.01
display "Question 1: PASSED"
```

## Question 2: Create Binary Variable (2 points)

Create a binary variable `is_foreign` that equals 1 if the car is foreign and 0 otherwise.

<!-- grade_id: q2_answer, points: 0, grade: false, solution: true -->
```stata
* Create is_foreign variable
* YOUR CODE HERE
gen is_foreign = foreign
```

<!-- grade_id: q2_test, points: 2, grade: true, solution: false -->
```stata
* Test Part 1: Variable exists (1 point)
capture confirm variable is_foreign
assert _rc == 0
display "Part 1: Variable exists - PASSED"

* Test Part 2: Values are correct (1 point)  
quietly count if foreign == 1 & is_foreign != 1
assert r(N) == 0
quietly count if foreign == 0 & is_foreign != 0
assert r(N) == 0
display "Part 2: Values correct - PASSED"
```

## Question 3: Simple Regression (2 points)

Run a regression of `price` on `mpg` and store the coefficient on `mpg` in a scalar called `mpg_coef`.

<!-- grade_id: q3_answer, points: 0, grade: false, solution: true -->
```stata
* Run regression and store coefficient
* YOUR CODE HERE
regress price mpg
scalar mpg_coef = _b[mpg]
```

<!-- grade_id: q3_test, points: 2, grade: true, solution: false -->
```stata
* Test
capture scalar list mpg_coef
assert _rc == 0

quietly regress price mpg
assert abs(mpg_coef - _b[mpg]) < 0.001

display "Question 3: PASSED"
display "MPG coefficient: " mpg_coef
```